rule_files:
  - availability99.yaml
evaluation_interval: 1m
group_eval_order: []
tests:
    - interval: 1m
      input_series:
        - series: http_requests_total{job="foo", code="200"}
          values: '999+999x1439 1438646+86x119 1449879+999x7079 '
        - series: http_requests_total{job="foo", code="500"}
          values: '1+1x1439 1454+14x119 3121+1x7079 '
      alert_rule_test:
        - eval_time: 1d1h30m
          alertname: SLOHighErrorRate
          exp_alerts:
            - exp_labels:
                job: foo
                slogen_id: foo-availability
                slogen_spec: foo
                slogen_slo: availability
                severity: page
              exp_annotations:
                description: 2% error budget has been consumed within 1 hour
      promql_expr_test:
        - expr: round(job:slogen_error:ratio_rate5m, 1e-6)
          eval_time: 1h
          exp_samples:
            - labels: '{job="foo", slogen_id="foo-availability", slogen_slo="availability", slogen_spec="foo"}'
              value: 0.001
        - expr: round(job:slogen_error:ratio_rate1h, 1e-6)
          eval_time: 1h
          exp_samples:
            - labels: '{job="foo", slogen_id="foo-availability", slogen_slo="availability", slogen_spec="foo"}'
              value: 0.001
        - expr: round(job:slogen_error:ratio_rate6h, 1e-6)
          eval_time: 1h
          exp_samples:
            - labels: '{job="foo", slogen_id="foo-availability", slogen_slo="availability", slogen_spec="foo"}'
              value: 0.001
        - expr: round(job:slogen_error:ratio_rate3d, 1e-6)
          eval_time: 1h
          exp_samples:
            - labels: '{job="foo", slogen_id="foo-availability", slogen_slo="availability", slogen_spec="foo"}'
              value: 0.001
        - expr: round(job:slogen_error:ratio_rate4w, 1e-6)
          eval_time: 1h
          exp_samples:
            - labels: '{job="foo", slogen_id="foo-availability", slogen_slo="availability", slogen_spec="foo"}'
              value: 0.001
        - expr: round(job:slogen_error_budget:ratio_rate4w, 1e-6)
          eval_time: 1h
          exp_samples:
            - labels: '{job="foo", slogen_id="foo-availability", slogen_slo="availability", slogen_spec="foo"}'
              value: 0.9
        - expr: slogen_slo
          exp_samples:
            - labels: 'slogen_slo{slogen_id="foo-availability", slogen_slo="availability", slogen_spec="foo"}'
              value: 0.99
        - expr: round(job:slogen_error:ratio_rate5m, 1e-6)
          eval_time: 1d1h30m
          exp_samples:
            - labels: '{job="foo", slogen_id="foo-availability", slogen_slo="availability", slogen_spec="foo"}'
              value: 0.14
        - expr: round(job:slogen_error:ratio_rate1h, 1e-6)
          eval_time: 1d1h30m
          exp_samples:
            - labels: '{job="foo", slogen_id="foo-availability", slogen_slo="availability", slogen_spec="foo"}'
              value: 0.14

