rule_files: []
evaluation_interval: 15m
group_eval_order: []
tests:
    - interval: 15m
      input_series:
        - series: http_requests_total{job="foo", code="200"}
          values: '999+999x95 '
        - series: http_requests_total{job="foo", code="500"}
          values: '1+1x95 '
      alert_rule_test: []
      promql_expr_test: []
